---
title: 'Python for Data Science'
description: 'Enjoy your exercise as you grow from zero to hero '
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

# **Hello,Welcome to Day 4- 5 of 100 days of Ds code.**

We discussed on Python libraries for Data Science and we mentioned 10 of them.
in this exercise we will focus on the popular and most used data science libraries such as 
1. Pandas
2. Numpy
3. Matplotlib
4. Scapy

`@instructions`
Type the 4 most used Python libraries for Data Science

`@hint`
PNMS

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
Pandas
Numpy
Matplotlib
Scapy
```

`@sct`
```{python}

```

---

## Numpy

```yaml
type: NormalExercise
key: 1f89ed8bc9
xp: 100
```

Wow, **Congrats** ```
:)
``` you have successfully entered the society of ds heros with python as your sword and you got to use your intuition as your skills
In this phase you will write code like a World-class Data Scientist. 
Let's look at the first python library on our list(previous session)

`@instructions`
Import```Numpy```  and pass it into an alias
```np``` then use ```np.around```use it to round up the sum of 2.345 + 2.7845 into 2 decimal places.

`@hint`
Import numpy and print (first number+ 2nd number), then use np.around to round them the sum total into 2 decimal places

`@pre_exercise_code`
```{python}
import numpy as np
value = (2.345 + 2.7845)
np.around(value,2)

```

`@sample_code`
```{python}

```

`@solution`
```{python}
import numpy as np
value = (2.345 + 2.7845)
np.around(value,2)
```

`@sct`
```{python}

```

---

## Numpy Exercise

```yaml
type: MultipleChoiceExercise
key: 20ccf0e81e
xp: 50
```

This exercise will requires you to research online 

Research and write 30 Numpy code to perform a task

`@possible_answers`


`@hint`
Click [here](https://www.w3resource.com/python-exercises/numpy/index.php) to see hint

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
