---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

# **Hello,Welcome to Day 4- 5 of 100 days of Ds code.**

We discussed on Python libraries for Data Science and we mentioned 10 of them.
in this exercise we will focus on the popular and most used data science libraries such as 
1. Pandas
2. Numpy
3. Matplotlib
4. Scapy

`@instructions`
Type the 4 most used Python libraries for Data Science

`@hint`
PNMS

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
Pandas
Numpy
Matplotlib
Scapy
```

`@sct`
```{python}

```
